Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Characteristic equations 2

```python
import numpy as np
```

A real LTIC system with input $x(t)$ and output $y(t)$ is described by the following constant-coefficient linear differential equation.
$$ (D^2 + 4D + 4){y(t)} = D{x(t)}$$

What is the characteristic equation of the system?  Represent your answer a function `Q(lam)`

<!-- grade_id: cell-8ef4018763dc4d60, points: 0, grade: false, solution: true -->
```python
def Q(lam):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2b699f0badd299cb, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q(1.5), float)

```

What are the characteristic roots?  Represent your answer as a list `ans = [XX, XX]`.

<!-- grade_id: cell-f40d36ecad1ba5fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18b4d687c43eb741, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, list)

```

Assuming $y_\text{zir}(0^-) = 3$ and $y'_\text{zir}(0^-) = -4$ determine the systems' zero-input response $y_\text{zir}(t)$.  You will need to simplify this to real terms in an exam, so I suggest you practice that here.

Repesent your answer as a function `y_zir(t)`.

<!-- grade_id: cell-7ff6ffd450a7a878, points: 0, grade: false, solution: true -->
```python
def y_zir(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e4a3c3f474be3802, points: 2, grade: true, solution: false -->
```python
assert isinstance(y_zir(1), float)


```