Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Characteristic equations 1

```python
import numpy as np
```

A real LTIC system with input $x(t)$ and output $y(t)$ is described by the following constant-coefficient linear differential equation.
$$ (D^3 +9D){y(t)} = (2D^3 + 1){x(t)}$$

What is the characteristic equation of the system?  Represent your answer a function `Q(lam)`

<!-- grade_id: cell-8da290933514cdac, points: 0, grade: false, solution: true -->
```python
def Q(lam):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7fbbfb255e0a9d9c, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q(1.5), float)

```

Assuming $y_\text{zir}(0) = 4$, $y'_\text{zir}(0) = -18$, and $y''_\text{zir}(0) = 0$, determine the systems' zero-input response $y_\text{zir}(t)$.  You will need to simplify this to real terms in an exam, so I suggest you practice that here.

Repesent your answer as a function `y_zir(t)`.

<!-- grade_id: cell-c94add431399aaa9, points: 0, grade: false, solution: true -->
```python
def y_zir(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a6404353f4b667e6, points: 2, grade: true, solution: false -->
```python
assert isinstance(y_zir(1), float)

```