Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Energy

```python
import numpy as np
```

The signal below is defined as
$$
x(t)=
\begin{cases}
t & 0 \leq t < 1\\
0.5 + 0.5\cos(2\pi t) & 1 \leq t < 2\\
3-t & 2 \leq t < 3 \\
0 & \text{otherwise}
\end{cases}
$$

What is the energy of $y_1(t) = (1/3)x(2t)$?

Represent your answer as `E_y1 = XX`.

<!-- grade_id: cell-b29f74eba2fa40cd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d2815115b529aa4, points: 1, grade: true, solution: false -->
```python
assert isinstance(E_y1, float)

```

A periodic signal $y_2(t)$ is defined as
$$
y_2(t) = 
\begin{cases}
x(t) & 0 \leq t < 4\\
y_2(t+4) & \forall t>
\end{cases}.
$$
What is the power of $y_2(t)$?

Represent your answer as `P_y2 = XX`.

<!-- grade_id: cell-070df7e1591d6319, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4286dff6e40e1581, points: 1, grade: true, solution: false -->
```python
assert isinstance(P_y2, float)

```

What is the power of $y_3(t) = (1/3)y_2(2t)$?

Represent your aswer as `P_y3 = XX`.

<!-- grade_id: cell-e91f6d7abb3df992, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f112a2ae77eb2f99, points: 1, grade: true, solution: false -->
```python
assert isinstance(P_y3, float)

```