Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Integrals

```python
import numpy as np
```

Evaluate the following integrals.  Represent your answer as `ans = XX`.

$$\int^\infty_{-\infty} \delta(t) e^{-j\omega t} dt$$

<!-- grade_id: cell-35e2c603555667d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ad7f401aac337d35, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, (int, float))

```

$$\int^\infty_{-\infty} \delta(2t-3) \sin(\pi t) dt$$

<!-- grade_id: cell-ef876ed2546514b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a6a5ce1703e9f834, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, (int, float))


```

$$\int^\infty_{-\infty} \delta(t+3) e^{-t} dt$$

<!-- grade_id: cell-66a9b34936ee69f8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-52f0ef560d75f9ca, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, (int, float))

```

$$\int^\infty_{-\infty} (t^3 +4)\delta(1-t)dt$$

<!-- grade_id: cell-a37a1b45330ddd93, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-994d3f20c300be14, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, (int, float))

```

$$\int^\infty_{-\infty} \delta(\tau) \sin(t-\tau)d\tau$$

Represent your answer as a function `f_1(t)`.

<!-- grade_id: cell-b6cdcf4c67be7164, points: 0, grade: false, solution: true -->
```python
def f_1(t):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-f134a91eaa57e9ef, points: 1, grade: true, solution: false -->
```python
assert isinstance(f(0), (int, float))

```

$$\int^\infty_{-\infty} \sin(\tau) \delta(t-\tau)d\tau$$

Represent your answer as a function `f_2(t)`.

<!-- grade_id: cell-4f382716c67ef454, points: 0, grade: false, solution: true -->
```python
def f_2(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-37d02cae37a50e27, points: 1, grade: true, solution: false -->
```python
assert isinstance(f(0), (int, float))

```