Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Linearity

```python
import numpy as np
```

For the systems desdcribed by the following equations, with the input $x(t)$ and the output $y(t)$, determine which of the systems are linear and which are nonlinear.

Represent your answer as `ans = "linear"` or `ans = "nonlinear"`.

$$\frac{dy(t)}{dt} + 2y(t) = x^2(t)$$

<!-- grade_id: cell-0a5f7b8c3d95a0f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eb9a7b7ac8e45745, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$\frac{dy(t)}{dt} + 3ty(t) = t^2 x(t)$$

<!-- grade_id: cell-f72636573a773f3b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-47f42b9d33fa9e3f, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$5y(t) + 1 = x(t)$$

<!-- grade_id: cell-5c5761bb8f44eae8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-63e6bc6b897eb27d, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$\left(\frac{dy(t)}{dt}\right)^2 +2y(t) = x(t)$$

<!-- grade_id: cell-cd304d67eed9c17a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0dab5fdcb37eafe2, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$\frac{dy(t)}{dt} + \sin(t) y(t) = \frac{dx(t)}{dt} + 2x(t)$$

<!-- grade_id: cell-509c20e97664c277, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-29c42383d5b824ec, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$\frac{d^2y(t)}{dt^2} + 5 \frac{dy(t)}{dt} + y(t) = \frac{dx(t)}{dt} + 2x(t)$$

<!-- grade_id: cell-47615bcd2648470b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3fd58b375b2d475, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```

$$y(t) = \int^t_{-\infty} x(\tau)d\tau$$

<!-- grade_id: cell-bb4f2aa5b6be5024, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8418b7b3af278380, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["linear", "nonlinear"])

```