Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Signal sketching

```python
from matplotlib import pyplot as plt
import numpy as np
```

This class will ask questions where the solution is not a simple numerical answer.  Often, it will be in a functional form.  

Consider $f(t) = u(t-5) -  u(t-7)$?

If you were to represent this in Python, how would you do this?  There are myriad approaches, but any of them will do the homework.  The simplest might be:

```python
def f_simple(t):
    if t < 5 or t > 7:
        return 0.
    else:
        return 1.
```

Alternatively, you could do the following, which is more elegant, but you aren't expected to produce the most elegant Python code for these homework exercises.

```python
def f_complicated1(t):
    f = np.zeros_like(t)
    f[5 < t <= 7] = 1.
    return f
```

```python
def f_complicated2(t):
    f = np.zeros_like(t)
    f[(5 < t) & (t <= 7)] = 1.
    return f
```

We see if we plot both you will get the same thing!

```python
t = np.linspace(3,8,1000)

plt.figure()
plt.plot(t, [f_simple(x) for x in t], label='simple')
plt.plot(t, f_complicated2(t), label='complicated')
plt.legend()
plt.xlabel('time')
plt.ylabel('amplitude')
plt.title('simple vs. complicated')
plt.grid()
plt.show()

```

Your turn!  Write a function `f_1(t)` for the expression $f_1(t) = t^2[u(t-1) - u(t-2)]$.

<!-- grade_id: cell-7491e846c95fe27e, points: 0, grade: false, solution: true -->
```python
def f_1(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-490de5d588c74818, points: 1, grade: true, solution: false -->
```python
assert isinstance(f_1(1.5), (int, float))

```

Let's try another.  Write a function `f_2(t)` for the expression $f_2(t) = (t-4) [u(t-2) - u(t-4)]$

<!-- grade_id: cell-1e03dea9e8f17200, points: 0, grade: false, solution: true -->
```python
def f_2(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-eadef12256f366ab, points: 1, grade: true, solution: false -->
```python
assert isinstance(f_2(1.5), (int, float))

```

Let's plot these to make sure they make sense intuitively.

```python
t = np.linspace(0,5,200)
f_1_array = [f_1(t_val) for t_val in t]
f_2_array = [f_2(t_val) for t_val in t]

plt.figure()
plt.plot(t, f_1_array, label="$f_1(t)$")
plt.plot(t, f_2_array, label="$f_2(t)$")
plt.legend()
plt.xlabel('time')
plt.ylabel('amplitude')
plt.grid()

```