Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Time invariance

```python
import numpy as np
```

For the systems described by the following equations, with the input $x(t)$ and the output $y(t)$, determine which are time-invariant and which are not.

Represent your answer as `ans = "TI"` or `ans = "not_TI"`.

$$y(t) = x(t-2)$$

<!-- grade_id: cell-25cdf124828e51cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ff61de3cede52942, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["TI", "not_TI"])

```

$$y(t) = \log(x(t))$$

<!-- grade_id: cell-7157b01ec651da14, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-04b9a783bd5f2185, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["TI", "not_TI"])

```

$$y(t) = t^2$$

<!-- grade_id: cell-037873b478f13204, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3ece75a0a850c526, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["TI", "not_TI"])

```

$$y(t) = \int^5_{-5} x(\tau) d\tau$$

<!-- grade_id: cell-c5c94f2d91470356, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5c7fadbc6276a9d7, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["TI", "not_TI"])

```

$$y(t) = t x(t) $$

<!-- grade_id: cell-5538e3f478e89fcd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5d5556ded98f0ae5, points: 1, grade: true, solution: false -->
```python
assert isinstance(ans, ["TI", "not_TI"])

```